![logo](nai-logo.png)

### Instructions

1. Make sure you are using a version of notebook greater than v.3. If you installed Anaconda with python 3 - this is likely to be fine. The next piece of code will check if you have the right version.
2. The notebook has both some open test cases that you can use to test the functionality of your code - however it will be run on another set of test cases that you can't from which marks will be awarded. So passing all the tests in this notebook is not a guarantee that you have done things correctly - though its highly probable.
3. Also make sure you submit a notebook that doesn't return any errors. One way to ensure this is to run all the cells before you submit the notebook.
4. When you are done create a zip file of your notebook and upload that
5. For each cell where you see "YOUR CODE HERE" delete the return notImplemented statement when you write your code there - don't leave it in the notebook.
6. Once you are done, you are done.

```python

```

## Machine Learning primer

This notebook takes you through how to load up Sklearn for doing Machine Learning. 

```python
from nose.tools import assert_equal, assert_greater, assert_less, assert_greater_equal
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

```python
# Import the good stuff
import pandas as pd
import numpy as np
from sklearn.linear_model import LogisticRegression
from sklearn.svm import SVC
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
```

#### Question 1

Read in the data from the file `multiclassdata.csv` using pandas

<!-- grade_id: cell-c523c9b97ee8ea91, points: 0, grade: false, solution: true -->
```python
def read_data():
    # Read in the multiclass dataset and output a numpy array of the data (without the column names)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1f10f049e89805fa, points: 2, grade: true, solution: false -->
```python
data = read_data()
assert_equal(isinstance(data, np.ndarray), True)
```

#### Question 2

Split the data into the features `X` and the labels `y`. The labels are last column of the data

<!-- grade_id: cell-c50a35e3a8d21d9d, points: 0, grade: false, solution: true -->
```python
def split_Xy(data):
    # Inputs a numpy array and outputs two arrays, X: the feature vectors and y: the labels
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f6226b45df4b15e6, points: 2, grade: true, solution: false -->
```python
X,y = split_Xy(data)
assert_equal(isinstance(X, np.ndarray), True)
assert_equal(X.shape[0], y.shape[0])
assert_equal(y.shape[0], 500)
```

#### Question 3

How many classes are represented in the label ?

<!-- grade_id: cell-23c37a2353f4a783, points: 0, grade: false, solution: true -->
```python
def num_classes(label):
    # Return the number of distinct classes in the label
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-15de29100c7723ef, points: 2, grade: true, solution: false -->
```python
nc = num_classes(y)
assert nc > 1
```

#### Question 4

Split the data into a test and training set both for X and y with 20% testset. Hint: you can use the Sklearn module train_test_split with a random_state of 10

<!-- grade_id: cell-93bfab55712fcee9, points: 0, grade: false, solution: true -->
```python
def split_train_test(X,y):
    # Split the data in Xtrain, Xtest, ytrain, ytest with a random_state of 10
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dee73705d9024973, points: 3, grade: true, solution: false -->
```python
X_train, X_test, y_train, y_test = split_train_test(X,y)
assert_equal(X_train.shape, (400,4))
assert_equal(y_test.shape, (100,))
```

#### Question 5

Train a Logistic regression classifer using Sklearn and return a prediction on the test set

<!-- grade_id: cell-f2f133bc4f91761d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8e27036ec9d8ec44, points: 4, grade: true, solution: false -->
```python
y_pred = logistic_regression(X_train, X_test, y_train)
assert_equal(y_pred.shape, (100,))
assert_greater(sum(y_pred), 100)
```

#### Question 6

Get the performance (accuracy) of your algorithm given ytest

<!-- grade_id: cell-d71553fa56907007, points: 0, grade: false, solution: true -->
```python
def get_accuracy(ypred, ytest):
    # Get Accuracy given ytest as a percentage
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-97d9aa38b652507f, points: 3, grade: true, solution: false -->
```python
acc = get_accuracy(y_pred, y_test)
assert_less(acc, 90)
```

#### Question 7

Retrain algorithm with Support Vector Classifier. Tune the parameters

<!-- grade_id: cell-a052f3664c53a74c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-539e705ce5951884, points: 4, grade: true, solution: false -->
```python
y_pred_svc = support_vector_classifier(X_train, X_test, y_train)
assert_equal(y_pred_svc.shape, (100,))
assert_greater(sum(y_pred_svc), 100)
```

#### Question 8

Get the accuracy of the SVC algorithm with default parameters

<!-- grade_id: cell-487b0540dab76013, points: 2, grade: true, solution: false -->
```python
ypredsvc = support_vector_classifier(X_train, X_test, y_train)
ac = get_accuracy(ypredsvc, y_test)
assert_less(ac, 87)
```

#### Question 9

Tune the parameters of the SVC algorithm to get a performance of greater than or equal to 92 %

<!-- grade_id: cell-5cbc6d2a0cd2db37, points: 3, grade: true, solution: false -->
```python
yypredsvc = support_vector_classifier(X_train, X_test, y_train, c=1000, kern='linear')
accc = get_accuracy(yypredsvc, y_test)
assert_greater_equal(accc, 92)
```

```python

```